<a href="https://colab.research.google.com/github/Meshack112/ASSIGNMENT-CIT-2228-ENG-219-058-2023/blob/main/Copy_of_Assignment_Control_Statements.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Exercise 02.1 (if-else)

Consider the following assessment criteria which map a score out of 100 to an
assessment grade:

| Grade            | Raw score  (/100)      |
| ---------------- | ---------------------- |
| Excellent        | $\ge 85$               |
| Very good        | $\ge 76.5$ and $< 85$  |
| Good             | $\ge 64$ and $< 76.5$  |
| Need improvement | $\ge 40$ and $< 64$    |
| Did you try?     | $< 40$                 |

Write a program that, given an a score, prints the appropriate grade. Print an error message if the input score is greater than 100 or less than zero.

<!-- grade_id: cell-71121f470be22e4a, points: 0, grade: false, solution: true -->
```python
# Score from user
score = 72
if score > 100 or score < 0:
    print("Error: Score must be between 0 and 100.")
elif score >= 85:
    print("Excellent")
elif score >= 76.5:
    print("Very good")
elif score >= 64:
    print("Good")
elif score >= 40:
    print("Need improvement")
else:
    print("Did you try?")
```

Output:
```
Good
```

## Exercise 02.2 (bisection)

Bisection is an iterative method for finding approximate roots of a function. Say we know that the function $f(x)$ has one root between $x_{0}$ and $x_{1}$ ($x_{0} < x_{1}$). We then:

- Evaluate $f$ at the midpoint $x_{\rm mid} = (x_0 + x_1)/2$, i.e. compute
   $f_{\rm mid} = f(x_{\rm mid})$
- Evaluate $f(x_0) \cdot f(x_{\rm mid})$

  - if $f(x_0) \cdot f(x_{\rm mid}) < 0$:

    $f$ must change sign somewhere between $x_0$ and $x_{\rm mid}$, hence the root must lie between
    $x_0$ and $x_{\rm mid}$, so set $x_1 = x_{\rm mid}$.
   
  - else:

    $f$ must change sign somewhere between $x_{\rm mid}$ and $x_1$, so set
    $x_0 = x_{\rm mid}$.

The above steps can be repeated a specified number of times, or until $|f_{\rm mid}|$
is below a tolerance, with $x_{\rm mid}$ being the approximate root.


### Task

The function

$$
f(x) = \frac{x^{5}}{10} + x^3 - 10x^2 + 4x + 7
$$


has one root in the range $0 < x < 2$.

1. Use the bisection method to find an approximate root $x_{r}$ using 20 iterations
   (use a `for` loop).
2. Use the bisection method to find an approximate root $x_{r}$ such that
   $\left| f(x_{r}) \right| < 1 \times 10^{-6}$ and report the number of iterations
   required (use a `while` loop).

Store the approximate root using the variable `x_mid`, and store $f(x_{\rm mid})$ using the variable `f`.

*Hint:* Use  `abs` to compute the absolute value of a number, e.g. `y = abs(x)` assigns the absolute value of `x` to `y`.

#### (1) Using a `for` loop.

```python
# Define the function
def f(x):
  return (x**5) / 10 + x**3 - 10 * x**2 + 4 * x + 7

# (1) Using a for loop
x0 = 0.0
x1 = 2.0

for n in range(20):
  x_mid = (x0 + x1) / 2
  f0 = f(x0)
  f_mid = f(x_mid)

  if f0 * f_mid < 0:
    x1 = x_mid
  else:
    x0 = x_mid

#print(n, x_mid, f_mid) #To see the output, run the code.


# (2) Using a while loop
x0 = 0.0
x1 = 2.0
tol = 1.0e-6
error = tol + 1.0
counter = 0

while error > tol:
  x_mid = (x0 + x1) / 2
  f0 = f(x0)
  f_mid = f(x_mid)
  error = abs(f_mid)

  if f0 * f_mid < 0:
    x1 = x_mid
  else:
    x0 = x_mid

  counter += 1

  if counter > 1000:
    print("Oops, iteration count is very large. Breaking out of while loop.")
    break

print(counter, x_mid, error)
```

Output:
```
24 1.1568354368209839 3.9180508970559913e-07
```

<!-- grade_id: cell-421824bbd4e61fae, points: 0, grade: true, solution: false -->
```python
## tests ##
import math
assert math.isclose(x_mid, 1.1568355560302734)
assert abs(f) < 1e-5
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-23-2fa7a19f1e12> in <cell line: 0>()
      1 ## tests ##
      2 import math
----> 3 assert math.isclose(x_mid, 1.1568355560302734)
      4 assert abs(f) < 1e-5

AssertionError:
```

#### (2) Using a `while` loop

Use the variable `counter` for the iteration number.

*Remember to guard against infinite loops.*

<!-- grade_id: cell-1747a3fb1dd1de48, points: 0, grade: false, solution: true -->
```python
import math
#Value of x (such that (1 + x) = 0.16)
x = -0.84

# Tolerance and initial error
tol = 1.0e-5
error = tol + 1.0

# Initialise approximation of function and counter
approx_f = 0.0
n = 0

# Loop until error satisfies tolerance, with a check to avoid an infinite loop
while error > tol and n < 1000:
    # Calculate the nth term of the series
    term = ((-1)**n * math.factorial(2 * n)) / (4**n * (math.factorial(n))**2) * x**n

    # Add the term to the approximation
    approx_f += term

    # Calculate the error
    error = abs(approx_f - (1 / math.sqrt(0.16 + 1))) # changed to calculate the error, which is the difference between the approximation and the actual value

    # Increment counter
    n += 1

#Print the results
print("\nThe error is:", error)
print("Number of terms in series:", n)
print("Approximation of 1/sqrt(0.16):", approx_f)
```

Output:
```
The error is: 1.5715233091147378
Number of terms in series: 1000
Approximation of 1/sqrt(0.16): 2.4999999999999973
```

<!-- grade_id: cell-ddff32f44df88388, points: 0, grade: true, solution: false -->
```python
## tests ##
assert counter == 24
assert abs(f) < 1.0e-6
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-22-26d7b8a156fa> in <cell line: 0>()
      1 ## tests ##
      2 assert counter == 24
----> 3 assert abs(f) < 1.0e-6

TypeError: bad operand type for abs(): 'function'
```

## Exercise 02.3 (series expansion)

For $|x| < 1$ the series:

$$
(1 + x)^{-1/2} = \sum_{n = 0}^{\infty} \frac{(-1)^n (2n)!}{4^n (n!)^2} x^n
$$

converges.

1. Using a `for` statement, approximate $1/\sqrt{0.16}$ using 30 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $1/\sqrt{0.16}$ to within $1 \times 10^{-5}$.

Store the absolute value of the error in the variable `error`.

### Hints

To compute the factorial, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
```
You only need `import math` once at the top of your program. Standard modules, like `math`, will be explained in a later

<!-- The power series expansion for the sine function is:

$$
\sin(x) = \sum_{n = 0}^{\infty} (-1)^n \frac{x^{2n +1}}{(2n+1)!}
$$

(See mathematics data book for a less compact version; this compact version is preferred here as it is simpler to program.)

1. Using a `for` statement, approximate $\sin(3\pi/2)$ using 15 terms in the series expansion and report the absolute error.

1. Using a `while` statement, compute how many terms in the series are required to approximate $\sin(3\pi/2)$ to within $1 \times 10^{-8}$.

Store the absolute value of the error in the variable `error`.

*Note:* Calculators and computers use iterative or series expansions to compute trigonometric functions, similar to the one above (although they use more efficient formulations than the above series).

### Hints

To compute the factorial and to get a good approximation of $\pi$, use the Python `math` module:
```python
import math
nfact = math.factorial(10)
pi = math.pi
```
You only need '`import math`' once at the top of your program. Standard modules, like `math`, will be explained in a later. If you want to test for angles for which sine is not simple, you can use
```python
a = 1.3
s = math.sin(a)
```    
to get an accurate computation of sine to check the error. -->

#### (1) Using a `for` loop

<!-- grade_id: cell-d6df4b428f969a5d, points: 0, grade: false, solution: true -->
```python
import math

# Value of x (such that (1 + x) = 0.16)
x = -0.84

# Initialise approximation of the function
approx_f = 0.0

# Loop for 30 terms
for n in range(30):
    # Calculate the nth term of the series
    term = ((-1)**n * math.factorial(2 * n)) / (4**n * (math.factorial(n))**2) * x**n

    # Add the term to the approximation
    approx_f += term

# Calculate the error
error = abs(approx_f - (1 / math.sqrt(0.16 + 1)))

print("The error is:", error)
print("Approximation of 1/sqrt(0.16):", approx_f)
```

Output:
```
The error is: 1.568331119541042
Approximation of 1/sqrt(0.16): 2.4968078104263016
```

<!-- grade_id: cell-8b269b717e92414d, points: 0, grade: true, solution: false -->
```python
## test ##
assert error < 1.0e-2
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-21-90d271ab5691> in <cell line: 0>()
      1 ## test ##
----> 2 assert error < 1.0e-2

AssertionError:
```

#### (2) Using a `while` loop

*Remember to guard against infinite loops.*

<!-- grade_id: cell-03df25d83df3674e, points: 0, grade: false, solution: true -->
```python
import math

# Value of x (such that (1 + x) = 0.16)
x = -0.84

# Tolerance and initial error
tol = 1.0e-5
error = tol + 1.0

# Initialise approximation of function and counter
approx_f = 0.0
n = 0

# Loop until error satisfies tolerance, with a check to avoid an infinite loop
while error > tol and n < 1000:
    # Calculate the nth term of the series
    term = ((-1)**n * math.factorial(2 * n)) / (4**n * (math.factorial(n))**2) * x**n

    # Add the term to the approximation
    approx_f += term

    # Calculate the error
    error = abs(approx_f - (1 / math.sqrt(0.16 + 1)))

    # Increment counter
    n += 1

# Print the results
print("\nThe error is:", error)
print("Number of terms in series:", n)
```

Output:
```
The error is: 1.5715233091147378
Number of terms in series: 1000
```

<!-- grade_id: cell-03df25d83df3674dfwe, points: 0, grade: true, solution: false -->
```python
## test ##
assert error <= 1.0e-5
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-25-1a61f82aef94> in <cell line: 0>()
      1 ## test ##
----> 2 assert error <= 1.0e-5

AssertionError:
```